# Review module

Throughout this pre-foundational program, you will complete several auto-graded mini-homeworks designed to test your understanding of the previous cases. This is the second of those mini-homeworks, covering case 0.2.

**Instructions**

The exercises will require you to write simple snippets of code. In order to complete this part, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).
4. Some functions may require several outputs (the docstrings tell you which ones). Make sure they are returned in the right order.

**How do functions work?**

You already saw some Python functions in the previous review module. Here's a reminder of some things you need to keep in mind:

1. Write your code as you would usually do, only this time write it inside the function just below the `### Your code here` comment.
2. Functions usually start with the `def` keyword and end with `return some_variable` . Don't modify any of these. If you just write your code as we explained above, you'll be good.
3. The code you write should be indented.
4. The docstring lets you know the name of the variable you will have to return. The `return` statement gives you this information as well. So if, for instance, the function says `return fancy_dict` you must make sure that the variable that you define as your result is also called `fancy_dict`.

Here are a couple of already-solved example questions:

### Example 1

Write code that takes a numerical value, multiplies it by itself, and then returns the result.

```python
def example_1(input_integer):
    """
    This is a docstring. Docstrings are written in plain English
    (they are not code), and are meant to help you remember
    what your function does. In this case, this function
    assigns `input_integer` to a variable, multiplies it by 2,
    and then returns the result. It takes an integer as input.
    
    Note: Hereinafter, we will be using the symbol ` to enclose words
    that are variable names, like we just did with `input_integer`
    """
    result = input_integer * input_integer
    return result
```

You can check that your function works by calling it like this:

```python
example_1(4)
```

Output:
```
16
```

### Example 2

Write code that takes any integer value, adds it to its own value divided by 10, and returns the result.

```python
def example_2(input_number):
    """
    Take an integer, add it to its own value divided by 10,
    and return the result.
    
    Arguments: (Arguments are the inputs you pass to your function)
    `input_number`: An integer
    """
    result = input_number + (input_number/10)
    return result
```

And checking:

```python
example_2(100)
```

Output:
```
110.0
```

Alright? Here come the coding questions!

### Exercise 1

Write code that finds the product of 4 numbers and returns it.

<!-- grade_id: cell-6da7fa29a659afd3, points: 0, grade: false, solution: true -->
```python
def multiply_these_numbers(a,b,c,d):
    """Multiplies `a`, `b`, `c`, and `d`
    and returns the product as the `result` variable
    """
    # YOUR CODE HERE
    result = a*b*c*d
    return result
```

### Exercise 2

Write code that adds (concatenates) 3 different strings together and returns the result.

<!-- grade_id: cell-b74d0ff90247b40b, points: 0, grade: false, solution: true -->
```python
def add_these_strings(string1, string2, string3):
    """
    Concatenates `string1`, `string2`, and `string3`.
    
    Since the arguments are strings, they must be surrounded
    by quotation marks when you pass them as inputs. So, for
    instance, you would call (execute) this function like this:
    
    add_these_strings("text one", "text two", "text three")
    
    Your function shouldn't add unnecessary spaces. The output of the
    previous call should then be
    
    "text onetext twotext three"
    
    and NOT
    
    "text one text two text three"
    """
    # YOUR CODE HERE
    result = string1+string2+string3
    return result
```

### Exercise 3

Write code that takes a number originally formatted as a string, divides it by the length of the original string, and returns the result.

<!-- grade_id: cell-466de4b7eb3a9f70, points: 0, grade: false, solution: true -->
```python
def number_gymnastics(number_as_string):
    """
    Take a number originally formatted as a string,
    then divide it by the length of the original string.
    
    Arguments:
    `number_as_string`: A number, only that it is formatted as a string.
    For instance: Instead of passing 15, you pass "15"
    (surrounded by quotation marks)
    
    Output:
    `result`: A number.
    
    Example:
    If you call number_gymnastics("14"),
    the output should be the number 7.0, because
    14 divided by 2 (the lenght of the "14" string) is equal
    to the number 7.0
    """
    # YOUR CODE HERE
    
    result = (int(number_as_string))/len(number_as_string)
    return result
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-e3ff00e471bec774, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert multiply_these_numbers(1,2,3,4) == 24, "Ex. 1 - Please check your arithmetic! Remember that to multiply numbers in Python, you use the symbol *"
assert multiply_these_numbers(2,3,4,5) == 120, "Ex. 1 - Please check your arithmetic! Remember that to multiply numbers in Python, you use the symbol *"
print("Exercise 1 looks correct!")
```

Output:
```
Exercise 1 looks correct!
```

<!-- grade_id: cell-c0480287df6d8dd5, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert add_these_strings("data", "is", "cool") == "dataiscool", "Ex. 2 - We tested your function, and the output isn't quite right yet. Remember that to concatenate two strings together, you can use the + symbol."
assert add_these_strings("I ", "need ", "space") == "I need space", "Ex. 2 - We tested your function, and the output isn't quite right yet. Remember that to concatenate two strings together, you can use the + symbol."
print("Exercise 2 looks correct!")
```

Output:
```
Exercise 2 looks correct!
```

<!-- grade_id: cell-bbb0af688a1c570a, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert number_gymnastics("270") == 90, "Ex. 3 - Not quite there yet! Remember, first transform the string input into a number (review case 0.2 if you need a refresher on how to do this), and then divide it by the length of the string input (again, check 0.2 to see how to find the lenght)"
assert number_gymnastics("5") == 5, "Ex. 3 - Not quite there yet! Remember, first transform the string input into a number (review case 0.2 if you need a refresher on how to do this), and then divide it by the length of the string input (again, check 0.2 to see how to find the lenght)"
print("Exercise 3 looks correct!")
```

Output:
```
Exercise 3 looks correct!
```