# Review module

## Part 1

For this first part, please uncomment the letter of the choice which you think is correct.

### Exercise 1

The first thing you should do when you are planning on using a new package or function from a package is to...

A. read the documentation

B. import it into your code

C. ask your friend who knows about the package to help

D. try to find an alternative way of solving your problem; packages should generally be avoided unless absolutely necessary

<!-- grade_id: cell-cac5f55b27c1d040, points: 0, grade: false, solution: true -->
```python
def ex_1():
    """Returns the option chosen by the fellow.
    """
    answer = "A"
    #answer = "B"
    #answer = "C"
    #answer = "D"
    
    # YOUR CODE HERE
    
    return answer
```

### Exercise 2

After reading the documentation, the next two steps you should take (in order) when trying to fix something you are stuck on are...

A. going to Stack Overflow and then reaching out to others

B. asking your TA and then asking your peers

C. asking your peers and then going to Stack Overflow

D. asking your co-workers and then asking your boss

<!-- grade_id: cell-eaf3303f896acbcf, points: 0, grade: false, solution: true -->
```python
def ex_2():
    """Returns the option chosen by the fellow.
    """
    answer = "A"
    #answer = "B"
    #answer = "C"
    #answer = "D"
    
    # YOUR CODE HERE
    
    return answer
```

## Part 2

**Instructions**

For this second part, the exercises will require you to write simple snippets of code. In order to complete this part, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).
4. Some functions may require several outputs (the docstrings tell you which ones). Make sure they are returned in the right order.

### Exercise 3

Write code that imports the `datetime` package from the standard library. Then, make it return the name of the day of the week of any date you pass to it (as a string, like "Monday"):

<!-- grade_id: cell-29a714e9affec8b5, points: 0, grade: false, solution: true -->
```python
def print_my_birthday_weekday(m,d,y):
    """Returns the day of the week of the date passed as an input
    
    Arguments:
    m: Month (as an integer)
    d: Day (as an integer)
    y: Year (as an integer)
    
    Outputs:
    day: The name of the day as a string, i.e., either
    'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday',
    'Saturday', or 'Sunday'
    
    Example:
    If you input 12,7,1991, `day` should be 'Saturday'
    """
    days = ['Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday']
    
    # YOUR CODE HERE
    import datetime
    day_of_week = datetime.date(y,m,d).weekday()
    day = days[day_of_week]
    
    return day
```

### Exercise 4

Write code that imports the `numpy` package and uses it to return the value of $\pi$ (pi).

<!-- grade_id: cell-ecce535f943efc5a, points: 0, grade: false, solution: true -->
```python
def return_pi():
    """Returns pi (a float)
    """
    
    # YOUR CODE HERE
    import numpy
    
    pi = numpy.pi
    return pi
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-05b4cc2a8cdf33e1, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert ord(ex_1()) == 65, "Ex. 1 - Think again! What is the very first thing you must do before you use a package?"
print("Exercise 1 seems correct!")
```

Output:
```
Exercise 1 seems correct!
```

<!-- grade_id: cell-c2a49990bbe407b6, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert int(ex_2().encode('utf-8').hex()) == 41, "Ex. 2 - Any of the options is good, really, but which one is the most efficient?"
print("Exercise 2 seems correct!")
```

Output:
```
Exercise 2 seems correct!
```

<!-- grade_id: cell-c6bcc541b23b37e1, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert print_my_birthday_weekday(12,7,1991) == 'Saturday', "Ex. 3 - Your function failed the test with Dec 7, 1991. What could've gone wrong? If you're stuck, remember you can always scour Stack Overflow to find solutions!"
assert print_my_birthday_weekday(12,9,1992) == 'Wednesday', "Ex. 3 - Your function failed the test with Dec 9, 1992. What could've gone wrong? If you're stuck, remember you can always scour Stack Overflow to find solutions!"
print("Exercise 3 seems correct!")
```

Output:
```
Exercise 3 seems correct!
```

<!-- grade_id: cell-7c1378554c9e3aaa, points: 1, grade: true, solution: false -->
```python
# Ex. 4
import inspect
s = inspect.getsource(return_pi)
assert ("import" in s) & ("numpy" in s) & ("pi" in s), "Ex. 4 - Did you import Numpy? Check case 0.7 if you need to review how to import libraries in Python! (Hint: Use the import keyword)"
assert abs(return_pi()*(5**2) - 78.53981633974483) < 0.01, "Ex. 4 - If you need help finding out how to compute pi using Numpy, why not ask on Stack Overflow?"
print("Exercise 4 seems correct!")
```

Output:
```
Exercise 4 seems correct!
```