# Review module

Throughout this pre-foundational program, you will complete several auto-graded mini-homeworks designed to test your understanding of the previous cases. This is the first of those mini-homeworks, covering case 0.1.

**Instructions**

For this review module, the exercises will require you to uncomment lines of code. In order to complete this part, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).

**How do functions work?**

You probably haven't seen Python functions before, but don't worry, we will be covering them in a future case. For now, the main things you need to know are:

1. Write your code as you would usually do, only this time write it inside the function just below the `### Your code here` comment.
2. Functions usually start with the `def` keyword and end with `return some_variable` . Don't modify any of these. If you just write your code as we explained above, you'll be good.
3. The code you write should be indented.
4. The docstring lets you know the name of the variable you will have to return. The `return` statement gives you this information as well. So if, for instance, the function says `return fancy_dict` you must make sure that the variable that you define as your result is also called `fancy_dict`.

Here are a couple of already-solved example questions:

### Example 1

What is the coolest animal in the world? (Remember, you need to uncomment the line that corresponds to the correct answer - you uncomment by removing the `#` symbol).

* A. Wombat
* B. Red panda.
* C. Hyrax.
* D. Meerkat.

```python
def coolest_animal():
    """
    This is a docstring. Docstrings are written in plain English
    (they are not code), and are meant to help you remember
    what your function does. In this case, this function
    assigns a letter to the variable `answer`,
    and then returns it. The letter corresponds
    to the coolest animal in the world.
    
    Note: Hereinafter, we will be using the symbol ` to enclose words
    that are variable names, like we just did with `answer`
    """
    #answer = "A"
    #answer = "B"
    #answer = "C"
    answer = "D"
    
    return answer
```

You can check that your function works by calling it like this:

```python
coolest_animal()
```

Output:
```
'D'
```

### Example 2

Why did the chicken cross the road? (Remember, you need to uncomment the line that corresponds to the correct answer - you uncomment by removing the `#` symbol).

* A. To attend an astrophysics conference.
* B. To chat with a cool meerkat.
* C. To get to the other side.
* D. To have its hair done.

```python
def chicken_road():
    """
    Assign a letter to the variable `answer`,
    and then return it. The letter corresponds
    to the reason why the chicken crossed the road.
    """
    #answer = "A"
    answer = "B"
    #answer = "C"
    #answer = "D"
    
    return answer
```

And checking:

```python
chicken_road()
```

Output:
```
'B'
```

In the questions below, you'll see a couple of lines that look like this:

~~~python
# YOUR CODE HERE
raise NotImplementedError() # Remove this line when you enter your solution
~~~

These lines are there to help us keep track of which exercises have been answered and which haven't. For the exercises in this review module, you should simply remove them and proceed to uncommenting the lines just as we explained above.

Alright? Here come the questions!

### Exercise 1

In computer science, Python is a(n)...

* A. IDE
* B. function
* C. programming language
* D. type of snake

<!-- grade_id: cell-19572b3e477e143f, points: 0, grade: false, solution: true -->
```python
def ex_1():
    """
    Assign a letter to the variable `answer`,
    and then return it.
    """
    
    # YOUR CODE HERE
    
    
    #answer = "A"
    #answer = "B"
    answer = "C"
    #answer = "D"
    
    return answer
```

### Exercise 2

The 3 disciplines necessary to become a great data professional are...

* A. Machine Learning, AI, Data Science

* B. Math & Statistics, Computer Science, Business Expertise

* C. Math & Statistics, Computer Programming, Data Science

* D. Math & Statistics, Python, Business Expertise

<!-- grade_id: cell-9f31b59057e28bc8, points: 0, grade: false, solution: true -->
```python
def ex_2():
    """
    Assign a letter to the variable `answer`,
    and then return it.
    """
    
    # YOUR CODE HERE
    
#    answer = "A"
    answer = "B"
#     answer = "C"
    #answer = "D"
    
    return answer
```

### Exercise 3

The `cd` command lets you...

* A. Rename folders
* B. See the contents of a folder
* C. Create a new folder
* D. Navigate around your computer's folders

<!-- grade_id: cell-b34688694bf03284, points: 0, grade: false, solution: true -->
```python
def ex_3():
    """
    Assign a letter to the variable `answer`,
    and then return it.
    """
    
    # YOUR CODE HERE
    
    #answer = "A"
    #answer = "B"
    #answer = "C"
    answer = "D"
    
    return answer
```

### Exercise 4

The `ls` command allows you to...

* A. see the name of the current file
* B. look into the contents of a file
* C. see the contents of a folder
* D. list the size of the current file

<!-- grade_id: cell-a7ef3f10e0b79b90, points: 0, grade: false, solution: true -->
```python
def ex_4():
    """
    Assign a letter to the variable `answer`,
    and then return it.
    """
    
    # YOUR CODE HERE
    
    #answer = "A"
#     answer = "B"
    answer = "C"
    #answer = "D"
    
    return answer
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-dfce3810f2c638eb, points: 1, grade: true, solution: false -->
```python
# Ex. 1
assert ord(ex_1()) == 67, "Ex. 1 - Think again! Check case 0.1 for a refresher!"
print("Exercise 1 seems correct!")
```

Output:
```
Exercise 1 seems correct!
```

<!-- grade_id: cell-f0e6846c9b920e1e, points: 1, grade: true, solution: false -->
```python
# Ex. 2
assert ord(ex_2())*10 == 660, "Ex. 2 - All these subjects are important, but which are the most important disciplines? See case 0.1 if you need a refresher!"
print("Exercise 2 seems correct!")
```

Output:
```
Exercise 2 seems correct!
```

<!-- grade_id: cell-71380078e8854049, points: 1, grade: true, solution: false -->
```python
# Ex. 3
assert ord(ex_3())*11 == 748, "Ex. 3 - This is the wrong answer. You can check the cd command's documentation - simply open your Jupyter terminal as explained in case 0.1 and type cd --help"
print("Exercise 3 seems correct!")
```

Output:
```
Exercise 3 seems correct!
```

<!-- grade_id: cell-1c9d233c0f481434, points: 1, grade: true, solution: false -->
```python
# Ex. 4
assert int(ex_4().encode('utf-8').hex()) == 43, "Ex. 4 - This is the wrong answer. You can check the ls command's documentation - simply open your Jupyter terminal as explained in case 0.1 and type ls --help"
print("Exercise 4 seems correct!")
```

Output:
```
Exercise 4 seems correct!
```