# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).
4. Some functions may require several outputs (the docstrings tell you which ones). Make sure they are returned in the right order.

**How do functions work?**

You probably haven't seen Python functions before, but don't worry, we will be covering them in a future case. For now, the main things you need to know are:

1. Write your code as you would usually do, only this time write it inside the function just below the `### Your code here` comment.
1. Functions usually start with the `def` keyword and end with `return some_variable` . Don't modify any of these. If you just write your code as we explained above, you'll be good.
1. The code you write should be indented.
1. The docstring lets you know the name of the variable you will have to return. The `return` statement gives you this information as well. So if, for instance, the function says `return fancy_dict` you must make sure that the variable that you define as your result is also called `fancy_dict`.

## Exercise 1

Make the following checklist into a dictionary. The systems marked with "yes" are proprietary, while those with "no" are open source. The names of the operating systems should be the keys (copy them exactly). Use [booleans](https://www.w3schools.com/python/python_booleans.asp) for the values.

> **Windows**: Yes; **Fire OS:** Yes; **macOS:** Yes; **Linux:** No; **Android:** No; **MS-DOS:** Yes.



<!-- grade_id: cell-cf8a32275814aaf9, points: 0, grade: false, solution: true -->
```python
def create_dictionary():
    """
    Create a dictionary.
    
    Arguments:
    None.
    
    Output:
    `my_dict`: A dictionary.
    
    """
    
    # YOUR CODE HERE
    my_dict={
        "Windows": True,
        "Fire OS": True,
        "macOS": True,
        "Linux": False,
        "Android": False,
        "MS-DOS": True        
    }
    
    
    return my_dict
```

## Exercise 2

In the function below you will find a list definition (`list_dup`). Write code to:

1. Remove the duplicate values.
2. [Sort](https://www.programiz.com/python-programming/methods/list/sort#) the items of the list in ascending alphabetical order.

The resulting list should be called `list_clean`.

<!-- grade_id: cell-fe7768571ff8d9de, points: 0, grade: false, solution: true -->
```python
def remove_dups_and_sort():
    """
    Remove the duplicates in the list below and then sort it
    alphabetically.
    
    Arguments:
    None.
    
    Outputs:
    `list_clean`: A sorted Python list with no duplicates.
    """
    
    list_dup = ["Defeat", "Defeat", "Turn of phrase", "Bob", "Cannot",
                "Eddy's", "Ana", "Bob", "Turn of phrase", "And", "Fantastic",
                "Bob", "Turn of phrase", "Cannot", "Defeat", "Fantastic", "Eddy's"]
    
    # YOUR CODE HERE
    unique_set=set(list_dup)
    unique_list=list(unique_set)
    
    list_clean = (sorted(unique_list))
    
    return list_clean
```

## Exercise 3

Code these two expressions as sets:

$$ A = \{1,2,3\}$$

$$ B = \{2,3,4,5\}$$

Find their intersection.

**Hint:** Remember the Venn diagram of an intersection:

![Venn diagram of an intersection](data/images/intersection_two.png)

<!-- grade_id: cell-de3938e1c0457316, points: 0, grade: false, solution: true -->
```python
def intersection_two_sets():
    """
    Return the intersection of sets A and B.
    
    Arguments:
    None.
    
    Outputs:
    `inter` = A Python set.
    """
    
    # YOUR CODE HERE
    A = {1, 2, 3}
    B = {2, 3, 4, 5}
    inter = A & B
    return inter
```

## Exercise 4

Do the same as in Exercise 3, but instead of finding the intersection of $A$ and $B$, find their union.

**Hint:** Remember the Venn diagram of a union:

![Venn diagram of an intersection](data/images/union.png)

<!-- grade_id: cell-9ab79188a11089ba, points: 0, grade: false, solution: true -->
```python
def union_two_sets():
    """
    Return the union of sets A and B.
    
    Arguments:
    None.
    
    Outputs:
    `uni` = A Python set.
    """
    
    # YOUR CODE HERE
    A = {1, 2, 3}
    B = {2, 3, 4, 5}
    uni = A | B
    return uni
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-cddcdf1ee0f8427d, points: 1, grade: true, solution: false -->
```python
# Ex 1
assert type(create_dictionary()) == type({"key":"value"}), "Ex. 1 - Your output is not a dictionary! Check case 0.3 if you need to review the syntax to create a dictionary in Python."
assert ['Android', 'Fire OS', 'Linux', 'MS-DOS', 'Windows', 'macOS'] == sorted(list(create_dictionary())), "Ex. 1 - Are you sure your dictionary has all the correct keys? Review the last part of case 0.3 to review how to create dictionaries in Python."
assert create_dictionary()["Windows"] == True, "Ex. 1 - Check your results! Windows is not open-source!"
assert create_dictionary()["macOS"] == True, "Ex. 1 - Check your results! macOS is not open-source!"
assert create_dictionary()["Fire OS"] == True, "Ex. 1 - Check your results! Fire OS is not open-source!"
assert create_dictionary()["MS-DOS"] == True, "Ex. 1 - Check your results! MS-DOS is not open-source!"
assert create_dictionary()["Android"] == False, "Ex. 1 - Check your results! Android is open-source!"
assert create_dictionary()["Linux"] == False, "Ex. 1 - Check your results! Linux is open-source!"
print("Exercise 1 looks correct!")
```

Output:
```
Exercise 1 looks correct!
```

<!-- grade_id: cell-6dd78b37092571c6, points: 1, grade: true, solution: false -->
```python
# Ex 2
assert type(remove_dups_and_sort()) == type([1,2,3]), "Ex. 2 - Are you sure your output is of type 'list'? You can check with type(). Remember that to convert a set into a list, you can use this code: list(my_set)"
assert not len(remove_dups_and_sort()) > 8, "Ex. 2 - Are you sure you removed the duplicates from your list? See the cheat sheet in case 0.3 to learn how to easily remove them!"
assert not len(remove_dups_and_sort()) < 8, "Ex. 2 - It appears you removed too many elements from your list. See the cheat sheet in case 0.3 to learn how to easily remove duplicates!"
assert remove_dups_and_sort() == ['Ana', 'And', 'Bob', 'Cannot', 'Defeat', "Eddy's", 'Fantastic', 'Turn of phrase'], "Ex. 2 - Your list isn't sorted or has duplicates! See the cheat sheet in case 0.3 to learn how to easily remove duplicates!"
print("Exercise 2 looks correct!")
```

Output:
```
Exercise 2 looks correct!
```

<!-- grade_id: cell-f7b4f6072e6f9ae6, points: 1, grade: true, solution: false -->
```python
# Ex 3
assert type(intersection_two_sets()) == type(set()), "Ex. 3 - Your result is not a set! You can check with type()."
assert len(list(intersection_two_sets())) == 2, "Ex. 3 - Are you sure that your intersection contains the right elements? Remember that the set intersection operator is `&` (see the cheat sheet in case 0.3 for more information)."
assert sum(list(intersection_two_sets())) == 5, "Ex. 3 - Are you sure that your intersection contains the right elements? Remember that the set intersection operator is `&` (see the cheat sheet in case 0.3 for more information)."
print("Exercise 3 looks correct!")
```

Output:
```
Exercise 3 looks correct!
```

<!-- grade_id: cell-b05bf733afe757c6, points: 1, grade: true, solution: false -->
```python
# Ex 4
assert type(union_two_sets()) == type(set()), "Ex. 4 - Your result is not a set. You can check with type()."
assert len(list(union_two_sets())) == 5, "Ex. 4 - Are you sure that your union contains the right elements? Remember that the set union operator is `|` (see the cheat sheet in case 0.3 for more information)."
assert sum(list(union_two_sets())) == 15, "Ex. 4 - Are you sure that your intersection contains the right elements? Remember that the set union operator is `|` (see the cheat sheet in case 0.3 for more information)."
print("Exercise 4 looks correct!")
```

Output:
```
Exercise 4 looks correct!
```