# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).
4. Some functions may require several outputs (the docstrings tell you which ones). Make sure they are returned in the right order.

**How do functions work?**

You probably haven't seen Python functions before, but don't worry, we will be covering them in a future case. For now, the main things you need to know are:

1. Write your code as you would usually do, only this time write it inside the function just below the `### Your code here` comment.
1. Functions usually start with the `def` keyword and end with `return some_variable` . Don't modify any of these. If you just write your code as we explained above, you'll be good.
1. The code you write should be indented.
1. The docstring lets you know the name of the variable you will have to return. The `return` statement gives you this information as well. So if, for instance, the function says `return fancy_dict` you must make sure that the variable that you define as your result is also called `fancy_dict`.

## Exercise 1

Below is a description of a fictional bureaucratic process. Make it a block of conditional statements.

> To qualify for benefits, you have to at least 18 years old and be either a citizen or a permanent resident. People who earn more than 50,000 dollars a year are not eligible.

Don't forget that your result variable should be named `eligible`.

**Hints:**

1. If you find it hard to follow the conditionals, you can always draw a decision tree!
2. To check if at least one of two conditions is true, you can use boolean logic.

<!-- grade_id: cell-ffff28d23d4fd7bb, points: 0, grade: false, solution: true -->
```python
def benefit_screening(applicant):
    """
    Determine if an applicant is eligible for a government benefit.
    
    Arguments:
    `applicant`: A Python list with all the relevant information.
    
    Output:
    `eligible`: A Python boolean.
    
    Example:
    Given the argument [22, False, True, 45000],
    the output variable `eligible` would have a value of `True`
    
    """
    
    age = applicant[0]
    citizen = applicant[1]
    resident = applicant[2]
    income = applicant[3]
       
    
    # YOUR CODE HERE
    if applicant[0] >= 18 and applicant[3] < 50000 and (applicant[1] == True or applicant[2] == True):
        eligible = True
    else:
        eligible = False
    return eligible
```

To test your function, you can run this code (you may modify the values in the list at will):

```python
applicant_no_1 = [17,False,True,5000] # Age, Citizen, Resident, Income
benefit_screening(applicant_no_1)
```

Output:
```
False
```

## Exercise 2

Complete the function below. Write the conditionals so that if the input `number` is even, the output variable `is_even` will be set to `True`. If the input is not even, `is_even` will be set to `False`.

**Hints:**

1. A number is even if when you divide it by 2, the remainder is equal to 0.
2. To find the remainder of a division in Python, you use the [**modulo** operator](https://www.freecodecamp.org/news/the-python-modulo-operator-what-does-the-symbol-mean-in-python-solved/). Like this: `a % b`.

<!-- grade_id: cell-17df40f0291a6753, points: 0, grade: false, solution: true -->
```python
def is_this_number_even(number):
    """
    Determine if `number` is even.
    
    Arguments:
    `number`: An integer.
    
    Outputs:
    `is_even`: A Python boolean.
    
    Example:
    Given the input 3, the output should be `False`
    """
    
    # YOUR CODE HERE
    if number % 2 == 0:
        is_even = True
    else:
        is_even = False
    
    return is_even
```

## Exercise 3

Rewrite this conditional block using `elif` statements:

~~~python
if my_dict["color"] == "brown":
    out = "Chocolate"
else:
    if my_dict["color"] == "red":
        out = "Strawberries"
    else:
        if my_dict["color"] == "green":
            out = "Avocado"
        else:
            out = "Unknown food"        
~~~

<!-- grade_id: cell-fba7cf3b3a00d32d, points: 0, grade: false, solution: true -->
```python
def rewrite_conditional(my_dict):
    """
    Rewrite a conditional block.
    
    Arguments:
    `my_dict`: A Python dictionary.
    
    Outputs:
    `out` = A Python string.
    
    Example:
    Given the input {"color":"red"},
    the output should be the string "Strawberries"
    """
    
    # YOUR CODE HERE
    if my_dict["color"] == "brown":
        out = "Chocolate"
    elif my_dict["color"] == "red":
        out = "Strawberries"
    elif my_dict["color"] == "green":
        out = "Avocado"
    else:
        out = "Unknown food"
    return out
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-e8b54fc508af9f32, points: 1, grade: true, solution: false -->
```python
# Ex 1
applicant = [15,True,True,1]
assert type(benefit_screening(applicant)) == type(True), "Ex. 1 - Your output doesn't seem to be a boolean! You can check with type()."
assert benefit_screening(applicant) == False, "Ex. 1 - One or more of your conditionals seem to be wrong. Perhaps your age conditional?"
applicant = [20,False,False,1]
assert benefit_screening(applicant) == False, "Ex. 1 - One or more of your conditionals seem to be wrong. Perhaps your citizen/resident conditional?"
applicant = [20,True,False,1]
assert benefit_screening(applicant) == True, "Ex. 1 - One or more of your conditionals seem to be wrong. Perhaps your citizen/resident conditional?"
applicant = [20,False,True,1]
assert benefit_screening(applicant) == True, "Ex. 1 - One or more of your conditionals seem to be wrong. Perhaps your citizen/resident conditional?"
applicant = [20,True,False,10000000]
assert benefit_screening(applicant) == False, "Ex. 1 - One or more of your conditionals seem to be wrong. Perhaps your income conditional?"
print("Exercise 1 looks correct!")
```

Output:
```
Exercise 1 looks correct!
```

<!-- grade_id: cell-39c8af6587e9ec0a, points: 1, grade: true, solution: false -->
```python
# Ex 2
assert is_this_number_even(157) == False, "Ex. 2 - Check your conditionals! We tested your function with 157, and it said it was even!"
assert is_this_number_even(5) == False, "Ex. 2 - Check your conditionals! We tested your function with 5, and it said it was even!"
assert is_this_number_even(-8) == True, "Ex. 2 - Check your conditionals! We tested your function with -8, and it said it was odd!"
print("Exercise 2 looks correct!")
```

Output:
```
Exercise 2 looks correct!
```

<!-- grade_id: cell-78a27010bc41fdeb, points: 1, grade: true, solution: false -->
```python
# Ex 3
import inspect
source_fun = inspect.getsource(rewrite_conditional)
assert "elif" in source_fun, "Ex. 3 - Did you use elif statements in your function?"
d = {"color":"brown"}
assert rewrite_conditional(d) == "Chocolate", "Ex. 3 - Check your conditionals. They aren't giving the right output for 'brown'"
d = {"color":"red"}
assert rewrite_conditional(d) == "Strawberries", "Ex. 3 - Check your conditionals. They aren't giving the right output for 'red'"
d = {"color":"green"}
assert rewrite_conditional(d) == "Avocado", "Ex. 3 - Check your conditionals. They aren't giving the right output for 'green'"
d = {"color":"purple"}
assert rewrite_conditional(d) == "Unknown food", "Ex. 3 - Check your conditionals. They aren't giving the right output for 'purple'"
print("Exercise 3 looks correct!")
```

Output:
```
Exercise 3 looks correct!
```