# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. Complete the functions provided to you in this notebook, but do **not** change the name of the function or the name(s) of the argument(s). If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!
3. Read the function docstrings carefully. They contain additional information about how the code should look (a [docstring](https://www.datacamp.com/community/tutorials/docstrings-python) is the stuff that comes between the triple quotes).
4. Some functions may require several outputs (the docstrings tell you which ones). Make sure they are returned in the right order.

## Exercise 1

Load the file that is in `data/time_series_covid19_confirmed_US.csv` and output its rows as a list.

<!-- grade_id: cell-3f6271a6e39561d5, points: 0, grade: false, solution: true -->
```python
def load_dataset():
    """
    Load the dataset that is in data/time_series_covid19_confirmed_US.csv
    
    Arguments:
    None.
    
    Output:
    `list_of_lines`: A list of the lines present in the file.
    
    """
    
    # YOUR CODE HERE
    file = open("data/time_series_covid19_confirmed_US.csv", "r")
    list_of_lines = file.readlines()
    return list_of_lines
```

## Exercise 2

Iterate through each of the lines of the file you've just loaded. If the line starts with `Alabama`, `Florida` or `Vermont`, store it in another list called `keep_lines`; if it doesn't, just ignore it.

**Hint:**

You can **split** a string using this syntax:

~~~python
my_string.split("character_to_split_by")
~~~

So, for instance,

~~~python
line.split(",")
~~~

Would split the line using a comma as the separator and output a list of strings. In the case of `time_series_covid19_confirmed_US.csv`, one of the strings in this list would correspond to the name of the state.

<!-- grade_id: cell-16e93f45c5f5b127, points: 0, grade: false, solution: true -->
```python
def keep_three_states():
    """
    Filters the dataset to keep only the data that corresponds to
    Alabama, Florida and Vermont.
    
    Arguments:
    None.
    
    Outputs:
    `keep_lines`: A list of the lines of the dataset that correspond to either
    of the three states.
    """
    
    dataset = load_dataset()
    
    keep_lines = []
    
    
    # YOUR CODE HERE
    for lines in dataset:
        state = lines.split(",")[0]
        if (state in "Alabama") or (state in "Florida") or (state in "Vermont"):
            keep_lines.append(lines)
    
    return keep_lines
```

## Exercise 3

Create a file called `afv.csv` and iterate through the lines produced by `keep_three_states()` to store them in the file. Save the file to disk (the path should be `data/afv.csv`).

<!-- grade_id: cell-d118393c487fd654, points: 0, grade: false, solution: true -->
```python
def save_afv():
    """
    Save the lines produced by `keep_three_states()`
    to a file called `afv.csv` in the `data` folder.
    
    Arguments:
    None.
    
    Outputs:
    None.
    """
    
    keep_list = keep_three_states()

    
    # YOUR CODE HERE
    keep_string = ''.join(map(str, keep_list))
    save_to_file = open("data/afv.csv", "w")
    line = keep_string
    save_to_file.write(line)
    save_to_file.close()

```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-6bcba8ed8779bb37, points: 1, grade: true, solution: false -->
```python
# Ex 1
assert len(load_dataset()) == 3341, "Ex. 1 - Please check your results. Your output should be a list of 3341 strings, and the first row should be the headers of the dataset. Have a look at the first section of case 0.6 for a refresher on how to load files into Python."
assert load_dataset()[0].split(",")[15] == '2/1/20', "Ex. 1 - Please check your results. Your output should be a list of 3341 strings, and the first row should be the headers of the dataset. Have a look at the first section of case 0.6 for a refresher on how to load files into Python."
print("Exercise 1 looks correct!")
```

Output:
```
Exercise 1 looks correct!
```

<!-- grade_id: cell-df896c1cc2afc536, points: 1, grade: true, solution: false -->
```python
# Ex 2
import pandas as pd
c = pd.Series(keep_three_states()).str.split(",", expand=True).groupby([0]).count()[1]
assert c[0] == 69, "Ex. 2 - Please check your results. Your dataset doesn't have the correct number of counties for one of the states. Have a look at case 0.4 if you need a refresher on conditional logic."
assert c[1] == 69, "Ex. 2 - Please check your results. Your dataset doesn't have the correct number of counties for one of the states. Have a look at case 0.4 if you need a refresher on conditional logic."
assert c[2] == 16, "Ex. 2 - Please check your results. Your dataset doesn't have the correct number of counties for one of the states. Have a look at case 0.4 if you need a refresher on conditional logic."
print("Exercise 2 looks correct!")
```

Output:
```
Exercise 2 looks correct!
```

<!-- grade_id: cell-b3474a68a1bbaa7e, points: 1, grade: true, solution: false -->
```python
# Ex 3
save_afv()
import os.path
from os import path
assert path.exists("data/afv.csv") == True, "Ex. 3 - Are you sure you created the file? It doesn't seem to exist! If the file does exist, check that it was saved with the right name in the right directory. Have a look at the last section of case 0.6 to brush up on saving files to disk in Python."
assert 154 == len(open("data/afv.csv", "r").readlines()), "Ex. 3 - Please check your results! Your file doesn't seem to have the right number of lines! Have a look at the last section of case 0.6 to brush up on saving files to disk in Python."
print("Exercise 3 looks correct!")
```

Output:
```
Exercise 3 looks correct!
```