# Review module

**Instructions**

In order to complete this review module, we recommend you follow these instructions:

1. When writing your code, do **not** use function names different to those provided to you in the prompts. If you do that, the autograder will fail and you will not receive any points.
2. Run all the function-definition cells before you run the testing cells. The functions must exist before they are graded!

## Exercise 1

Write a function called `cube_root` that takes a nonnegative number as its only argument. This function must find the cube root of the number and return it with a `return` statement.

**Note:** For this exercise, just assume the user will always input a nonnegative number - no need to include input validity checks.

**Hint:** Remember that for $X \geq 0$,

$$
\sqrt[3]X = X^{1/3}
$$

<!-- grade_id: cell-27de682534bc2137, points: 0, grade: false, solution: true -->
```python
# Please name your function cube_root - If you use another name, the test will fail!
# Write your function below (check case 0.5 to see how to define a function in Python)
# YOUR CODE HERE
def cube_root(num):
    num = num**(1/3)
    return num

```

## Exercise 2

Write an appropriate docstring for this function:

<!-- grade_id: cell-b75f1fddc7d0f9b1, points: 0, grade: false, solution: true -->
```python
def my_function(my_string, reps): # Don't change this line!
    
    # YOUR CODE HERE
    """
    Convert string to upper case and then repeat line as many times as assigned.
    
    Arguments:
    my_string: A line of letters, numbers or words that's given.
    reps: The amount of times the string needs to be repeated.
    
    Outputs:
    Repeated string in uppercase.
    """
    
    my_string = my_string.upper()
    repeated = my_string * reps
        
    return repeated
```

## Exercise 3

Write a function called `process_number` that takes two positive integers as arguments. This function should:

1. Find the cube root of the first integer (for this, we recommend you call `cube_root` inside this function), and then
2. take that cube root and elevate it to a power for which the exponent is the second number.

For instance, if the two integers are $x=27$ and $y=4$, then the output of this function should be

$$
(\sqrt[3]x)^y = (\sqrt[3]27)^4 = 3^4 = 81
$$

**Note:** For this exercise, just assume the user will always input positive numbers - no need to include input validity checks.

<!-- grade_id: cell-59e7e19785f29189, points: 0, grade: false, solution: true -->
```python
# Please name your function process_number - If you use another name, the test will fail!
# Write your function below (check case 0.5 to see how to define a function in Python)
# YOUR CODE HERE
def process_number(x, y):
    answer = (x**(1/3))**y
    return answer


```

## Exercise 4

Write a function called `split_number` that takes a float number as input and outputs two floats:

1. The first float is the "integer" part of the number. So, if the input was $25.7$, this first float should be $25.0$.
2. The second float is the "decimal" part of the number. So, if the input was $25.7$, this second float should be $7.0$ (*not* $0.7$).

**Hint:** For this, we suggest that you first convert the number into a string and then use `.split(".")` to separate this string into its components. Don't forget to return both outputs as **floats**.

<!-- grade_id: cell-e608ada7b2b43e72, points: 0, grade: false, solution: true -->
```python
def split_number(number):
    """
    Split `number` into its integer and decimal parts.
    
    Arguments:
    `number`: An integer or float
    
    Outputs:
    `integer_part`: A float.
    `decimal_part`: A float.
    
    Example:
    split_number(25.7) should produce 25.0 as the first output
    and 7.0 as the second output.
    """
    # YOUR CODE HERE
    num_to_string = str(number).split(".")
    integer_part = float(num_to_string[0])
    decimal_part = float(num_to_string[1])
    return integer_part, decimal_part
```

## Testing Cells

Run the below cells to check your answers. Make sure you run your solution cells first before running the cells below, otherwise you will get a `NameError` when checking your answers.

<!-- grade_id: cell-293c1c13c0baad1c, points: 1, grade: true, solution: false -->
```python
# Ex 1
try:
    cube_root(27)
except NameError:
    assert 3**3 == 9, "Ex. 1 - You haven't defined a function named cube_root! Remember to name your function cube_root and execute the corresponding cell before running this test."
else:
    assert round(cube_root(27)) == 3, "Ex. 1 - Incorrect! Remember that to elevate a number x to the y-th power, you use x**y in Python."
    assert round(cube_root(1000)) == 10, "Ex. 1 - Incorrect! Remember that to elevate a number x to the y-th power, you use x**y in Python."
    print("Exercise 1 looks correct!")
```

Output:
```
Exercise 1 looks correct!
```

<!-- grade_id: cell-303c96e4e07385c0, points: 1, grade: true, solution: false -->
```python
# Ex 2
import inspect
source_fun = inspect.getsource(my_function)
assert source_fun.count('"""') == 2, "Ex. 2 - Did you add a docstring to your function? Remember to use triple quotes both at the beginning and at the end of the docstring."
assert len(source_fun.split('"""')[1]) > 15, "Ex. 2 - Did you add a docstring to your function? Short docstrings are good, but sometimes we need more information!"
assert "argument" in source_fun.split('"""')[1].lower(), "Ex. 2 - Did you specify the arguments of the function in your docstring?"
assert "output" in source_fun.split('"""')[1].lower(), "Ex. 2 - Did you specify the output of the function in your docstring?"
print("Exercise 2 looks correct!")
```

Output:
```
Exercise 2 looks correct!
```

<!-- grade_id: cell-0a378773c1097e13, points: 1, grade: true, solution: false -->
```python
# Ex 3
try:
    process_number(27, 2)
except NameError:
    assert 3**3 == 9, "Ex. 3 - You haven't defined a function named process_number! Remember to name your function process_number and execute the corresponding cell before running this test."
else:
    assert round(process_number(27, 2)) == 9, "Ex. 3 - Your function is not producing the correct output. Please check your arithmetic!"
    assert round(process_number(64, 4)) == 256, "Ex. 3 - Your function is not producing the correct output. Please check your arithmetic!"
    print("Exercise 3 looks correct!")
```

Output:
```
Exercise 3 looks correct!
```

<!-- grade_id: cell-9003ad0a93088895, points: 1, grade: true, solution: false -->
```python
# Ex 4
assert split_number(27.2)[0] == 27, "Ex. 3 - Your function is not producing the correct output. Please check your arithmetic!"
assert split_number(64.4)[1] == 4, "Ex. 3 - Your function is not producing the correct output. Please check your arithmetic!"
assert split_number(12.0)[1] == 0, "Ex. 3 - Your function is not producing the correct output. Please check your arithmetic!"
print("Exercise 4 looks correct!")
```

Output:
```
Exercise 4 looks correct!
```